rule_files:
    - alerts.yml

evaluation_interval: 1m

tests:
    - interval: 1m
      input_series:
          - series: 'go_goroutines{job="prometheus", instance="localhost:9090"}'
            values: '10+10x2 30+20x5' # 10 20 30 30 50 70 90 110 130
          - series: 'go_goroutines{job="node_exporter", instance="localhost:9100"}'
            values: '10+10x7 10+30x4' # 10 20 30 40 50 60 70 80 10 40 70 100 130

      promql_expr_test:
          # Query returning no data.
          - expr: go_goroutines >= 50
            eval_time: 1m
            exp_samples:
                - labels: 'go_goroutines{job="prometheus",instance="localhost:9090"}'
                  value: 50
                - labels: 'go_goroutines{job="node_exporter",instance="localhost:9100"}'
                  value: 50
          # Missing sample.
          - expr: go_goroutines >= 50
            eval_time: 4m
            exp_samples:
                - labels: 'go_goroutines{job="prometheus",instance="localhost:9090"}'
                  value: 50
          # Wrong values.
          - expr: go_goroutines >= 50
            eval_time: 5m
            exp_samples:
                - labels: 'go_goroutines{job="prometheus",instance="localhost:9090"}'
                  value: 50
                - labels: 'go_goroutines{job="node_exporter",instance="localhost:9100"}'
                  value: 50

